---
title: '2. Python Lists'
description: ""
---

## 2.1. Lists ex.

```yaml
type: NormalExercise
key: a914db4710
xp: 100
```



`@instructions`
Create a list called myList with the following six items: 76, 92.3, “hello”, True, 4, 76. Do it with both append and with concatenation, one item at a time.
Starting with the list of the previous exercise, write Python statements to do the following:

1. Append “apple” and 76 to the list.
2. Insert the value “cat” at position 3.
3. Insert the value 99 at the start of the list.
4. Find the index of “hello”.
5. Count the number of 76s in the list.
6. Remove the first occurrence of 76 from the list.
7. Remove True from the list using pop and index.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: c4114ca86b
xp: 50
```

A list can contain only integer items.

`@possible_answers`
- True
- [False]

`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
