---
title: '6. Loading Data in Pandas'
description: ""
---

## 6.1. Car Park ex.

```yaml
type: NormalExercise
key: 08fb6ad600
xp: 100
```

You are a the owner of the legendary “XX” car park. Some investors came in to offer you a deal but they first want to know some information, which could be shown with the help of pandas and python. How lucky they are that you actually do know Pandas! But first, import the car data from your park ("https://raw.githubusercontent.com/guipsamora/pandas_exercises/master/05_Merge/Auto_MPG/cars2.csv"). 

`@instructions`
Now find the following:
1. How many cars weigh less than 2000kg?
2. What is the percentage of cars out of the whole park that have 8 cylinders?
3. What is the most expensive car?

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## 6.2. Baby Boom ex. 

```yaml
type: NormalExercise
key: 0d4239c7e6
xp: 100
```

You work at the department that counts and records the baby names in a given year. You are bored, but your boss wants you to identify some statistics. Your boss told you to do it by hand, but you know pandas and python, which would help you to get finished with this task. The source for the baby statistics: ('https://raw.githubusercontent.com/guipsamora/pandas_exercises/master/06_Stats/US_Baby_Names/US_Baby_Names_right.csv'). 

`@instructions`
Now find out the following: 

1. Are there more male or female names in the dataset?
2. What is the name with most occurrences?

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## 7.1. Plotting data

```yaml
type: NormalExercise
key: 3fac358f6a
xp: 100
```

Imagine yourself working for a company, hired to analyze the user growth rate of the company. For an analysis, you are required to visualize the data and show whether the growth is linear, exponential, etc. 

`@instructions`
Here is the sample data:
 
```
Year Number of users
2011    43
2012    124
2013    400
2014    682
2015    793
2016    899
2017    977
2018    1201
2019    1534
 
```



 How would you visualize it? Make sure you show the dates, label the axes, put a title, and make the graph beautiful!

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
