---
title: '10. Loops'
description: ""
---

## 10.1. ex.

```yaml
type: NormalExercise
key: 2b50d303a9
xp: 100
```



`@instructions`
Write a program that prints We like Python's turtles! 100 times.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## 10.2.

```yaml
type: NormalExercise
key: 42c59d43f2
xp: 100
```



`@instructions`
Write a Python program that accepts a word from the user and reverse it.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## 10.3.

```yaml
type: NormalExercise
key: 885ec5ddf4
xp: 100
```



`@instructions`
Write a Python program that accepts a string and calculate the number of digits and letters.
```
Sample Data : Python 3.2
Expected Output :
Letters 6 
Digits 2
```

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## 10.4.

```yaml
type: NormalExercise
key: 6bcb68a16a
xp: 100
```



`@instructions`
Write a Python program to create the multiplication table (from 1 to 10) of a number.
```
Expected Output:

Input a number: 6                                                       
6 x 1 = 6                                                               
6 x 2 = 12                                                              
6 x 3 = 18                                                              
6 x 4 = 24                                                              
6 x 5 = 30                                                              
6 x 6 = 36                                                              
6 x 7 = 42                                                              
6 x 8 = 48                                                              
6 x 9 = 54                                                              
6 x 10 = 60 
```

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
