---
title: '1. Python Basics'
description: 'After completing the ​Python Basics Assignment - please do the following exercises to check your understanding.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`
Take the sentence: "All work and no play makes Jack a dull boy". 
Store each word in a separate variable, then print out the sentence on one line using print.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
