---
title: 'Introduction to Matplotlib Exercises '
description: ""
---

## 7.1. Plotting data

```yaml
type: NormalExercise
key: bd365b4bcc
xp: 100
```

Imagine yourself working for a company, hired to analyze the user growth rate of the company. For an analysis, you are required to visualize the data and show whether the growth is linear, exponential, etc.

`@instructions`
Here is the sample data:
 
```
Year Number of users
2011    43
2012    124
2013    400
2014    682
2015    793
2016    899
2017    977
2018    1201
2019    1534
 
```



 How would you visualize it? Make sure you show the dates, label the axes, put a title, and make the graph beautiful!

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## 7.2. Heart graphing ex. <3

```yaml
type: NormalExercise
key: 1cc95c5cfd
xp: 100
```



`@instructions`
Let's imagine you want to show gratitude to those who took care of you. You're now looking for a heart image online to print it out and glue it to your card. However, you can't find anything that looks good. That said, your job is create a heart graph that is absolutely pretty.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
