---
title: '8. Dictionaries and Pandas'
description: ""
---

## TBD

```yaml
type: MultipleChoiceExercise
key: d24c684f4a
xp: 50
```

I am not sure if we need any ex. at all?
we may just delete this chapter if not needed

`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
