---
title: '13. ACCUMULATOR PATTERN'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: b4e726e6f0
xp: 100
```

TBD
https://interactivepython.org/courselib/static/thinkcspy/Functions/TheAccumulatorPattern.html 

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
