---
title: '8. Dictionaries and Pandas'
description: ""
---

## TBD

```yaml
type: MultipleChoiceExercise
key: d24c684f4a
xp: 50
```

Explore official Python Documentation - Dictionaries Topic[](https://docs.python.org/3/tutorial/datastructures.html#dictionaries). Use the console to experiment and #comments to brainstorm some ideas in what programs you can use dictionaries. 



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
