---
title: '5. Getting Started in Python'
description: ""
---

## 5.1. Pi estimation ex.

```yaml
type: NormalExercise
key: ca175201e4
xp: 100
```



`@instructions`
Search on the internet for a way to calculate an approximation for pi. There are many that use simple arithmetic. Write a program to compute the approximation and then print that value as well as the value of math.pi from the math module.

`@hint`
Some methods to try!
1) https://academo.org/demos/estimating-pi-monte-carlo/

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
